## Assignment 2

## Imports
It is important for all (or most) imports to go on the top of a notebook so that other users know which packages need to be installed. In projects that use Anaconda, it is also common to see a file named requirements.txt listing all the packages that one has to install.

First, let's import all the necessary modules using the import function. For this exercise we will continue to use Pandas, and will gain more experience using Numpy, Matplotlib, and Seaborn.

To learn more about these packages, you can read through the documentation:  
https://pandas.pydata.org/  
https://numpy.org/  
https://seaborn.pydata.org/  
https://matplotlib.org/

```python
# Let's import the relevant Python packages here
# Feel free to import any other packages for this project

#Data Wrangling
import pandas as pd
import numpy as np

#Sci Kit Learn
from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import PolynomialFeatures

#Stats Models
import statsmodels.formula.api as smf
import statsmodels.api as sm

#Metrics
from sklearn.metrics import mean_squared_error

#Plotting
import matplotlib.pyplot as plt

%matplotlib inline
```

## Data 1
First, we will use the Data1 dataset to generate polynomial regressions using scikitlearn.
This dataset contains 100 observations of points x and their corresponding response, y. The data
is divided into a training set (xtr, ytr) and a test set (xte, yte).

Let's load the Data1 dataset using pandas and check a couple of rows 

```python
Data1 = pd.read_csv("Data1.csv")
Data1.head()
```

Let's create a scatter plot of y_tr against x_tr 

```python
plt.plot(Data1['x_tr'],Data1['y_tr'],'o')
plt.xlabel('x')
plt.ylabel('y')
plt.title("Plot of training dataset")
plt.show()
```

```python
# Let's try to find the maximum y value in training data 

print("Max value of y in the training set: %s " % Data1.y_tr.max())
```

1. Now, make a scatter plot of y_te against x_te. Find the maximum value of y in testing data and save this in a variable called `max_y_te`. Notice the similarities and differences between the plots.

<!-- grade_id: cell-c62781c083043e86, points: 0, grade: false, solution: true -->
```python
# Write code for scatter plot like above

# your code here
raise NotImplementedError
```

<!-- grade_id: cell-17bfab400137c926, points: 0, grade: false, solution: true -->
```python
max_y_te = None # Store the maximum value of y in testing data 
print("Max value of y in the training set: %s " % max_y_te)

# your code here
raise NotImplementedError
```

<!-- grade_id: cell-032edd4d7fd0ad8, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

import math
assert math.isclose(max_y_te,0.865, rel_tol=0.001), "Max value of y in the testing set is wrong, please try again.."
```

### Observation:
Both the training set and test set range from 0 to approximately 1.0 on the x-axis. However, they differ slightly in range on the y-axis, as indicated below. Both sets have a negative polynomial relationship that seems to fit an odd-powered polynomial like $y=x^3$  or  $y=x^5$ . It seems, however, that the training set has a few end points that trend upward, while the test set has end points that shift mostly downward, which is bit concerning because we will only be training our model on the training set.

```python
Data1.describe()
```

2. Perform the following 
-  Generate the necessary features to fit polynomial regressions up to the 20th degree (up to and including the x^20 term) on the training data. 
    
-  Fit Linear regression model
-  Calculate mean squared error on train and test data

Hint: You will be fitting multi-variate linear regression models with polynomial features of x. Familiarize yourself with sklearn.preprocessing.PolynomialFeatures.

<!-- grade_id: cell-solution-4, points: 0, grade: false, solution: true -->
```python
x_tr = Data1[['x_tr']]
y_tr = Data1[['y_tr']]
max_degree = 21
mse_train_list = []
mse_test_list = []

## Fit polynomial regressions from 1 to 20 degrees 
for d in range(1,max_degree):
    
    polynomial_features = None # Hint: Use PolynomialFeatures() with degree d 
    
    poly = None # Hint: Fit transform on polynomialFeatures using 'x' train data
    
    poly_model = None # Hint: Use a model like LinearRegression and fit the model using 'poly' and 'y' train data
    
    # your code here
    raise NotImplementedError
    
    ## Predict training values using the fitted poly_model
    y_train_pred = poly_model.predict(poly)
    
    mse_train = None # Hint: Calculate the mean squared error between y_tr and y_train_pred
   
    # your code here
    raise NotImplementedError

    mse_train_list.append(mse_train) 

    y_test_pred = poly_model.predict(polynomial_features.fit_transform(Data1[['x_te']])) 
    mse_test = None # Hint: Calculate the mean squared error between y_te and y_test_pred
    
    # your code here
    raise NotImplementedError
    
    mse_test_list.append(mse_test)
    
    
plt.subplot(2,1,1)
plt.plot(range(1,max_degree), mse_train_list, '-r')
plt.ylabel("Training MSE")
plt.subplot(2,1,2)
plt.plot(range(1,max_degree), mse_test_list, '-b')
plt.xlabel("Flexibility (degree of polynomial)")
plt.ylabel("Test MSE")
plt.show()
```

<!-- grade_id: cell-032bedd4d7fd0ad8, points: 40, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert abs(min(mse_train_list) - 0.024) < 0.001, "Minimum value in mse_train_list is wrong. Please try again.."
assert abs(min(mse_test_list) - 0.032) < 0.001, "Minimum value in mse_test_list is wrong. Please try again.."
assert abs(mse_train_list[4] - 0.029) < 0.001, "One of the values in mse_train_list is wrong. Please try again.."
assert abs(mse_test_list[4] - 0.032) < 0.001, "One of the values in mse_test_list is wrong. Please try again.."

print("test passed!")
```

## Data 2 

Next, we will use the Ch3PartB dataset to observe the effects of collinearity using statsmodels.
This dataset contains 100 observations of points x1, x2, and y, the response variable.

```python
Data2 = pd.read_csv("Data2.csv").copy()
Data2.head()
```

```python
# Let's check the scatter plot for this data

plt.scatter(Data2.x1, Data2.x2)
plt.xlabel('x1')
plt.ylabel('x2')
plt.title('Relationship between x1 and x2')
plt.show()
```

3. Using the data, fit a least squares regression to predict y using x1 and x2.

<!-- grade_id: cell-solution-7, points: 0, grade: false, solution: true -->
```python
fit_both = None # Hint: Use smf ols to fit regression to predict y using x1 and x2

# your code here
raise NotImplementedError

fit_both.summary()

```

4. From the above summary
- Fill in the estimates of coefficients. 
- At a 95% confidence level, can you reject the null hypothesis H0 : β2 = 0? 

<!-- grade_id: cell-solution-8, points: 0, grade: false, solution: true -->
```python
# Save your final solutions in these variable

beta0 = None 
beta1 = None
beta2 = None
reject_hypothesis = None # Fill in this variable with either True or False for H0 : β2 = 0

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-8108a4e9b615b0f8, points: 30, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert abs(beta0 - 2.189) < 0.001, "Beta0 value is wrong. Please try again.."
assert abs(beta1 - 0.704) < 0.001, "Beta1 value is wrong. Please try again.."
assert abs(beta2 - 2.502) < 0.001, "Beta2 value is wrong. Please try again.."
assert reject_hypothesis == True, "Your reject_hypothesis decision in wrong. Please try again.. "
```

5. Now fit a least squares regression to predict y using only x1.

<!-- grade_id: cell-solution-10, points: 0, grade: false, solution: true -->
```python
fit_x1_only = None # Hint: Use smf ols to fit regression to predict y using only x1 

# your code here
raise NotImplementedError

fit_x1_only.summary()
```

6. At a 95% confidence level, can you reject the null hypothesis H0 : β1 = 0?

<!-- grade_id: cell-solution-108, points: 0, grade: false, solution: true -->
```python
reject_hypothesis = None  # Fill in this variable with either True or False for H0 : β1 = 0

# your code here
raise NotImplementedError
```

<!-- grade_id: cell-solution-118, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert reject_hypothesis == True, "Your reject_hypothesis decision in wrong. Please try again.. "
```

7. Now fit a least squares regression to predict y using only x2

<!-- grade_id: cell-solution-102, points: 0, grade: false, solution: true -->
```python
fit_x2_only = None # Hint: Use smf ols to fit regression to predict y using only x2 

# your code here
raise NotImplementedError

fit_x2_only.summary()
```

8. At a 95% confidence level, can you reject the null hypothesis H0 : β1 = 0?

<!-- grade_id: cell-solution-104, points: 0, grade: false, solution: true -->
```python
reject_hypothesis = None # Fill in this variable with either True or False for H0 : β1 = 0

# your code here
raise NotImplementedError
```

<!-- grade_id: cell-solution-114, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert reject_hypothesis == False, "Your reject_hypothesis decision in wrong. Please try again.. "
```