# Assignment 1: Exploratory Data Analysis

## Note
Make sure to remove the line **raise NotImplementedError** when you add your code.


## Imports
It is important for all (or most) imports to go on the top of a notebook so that other users know which packages need to be installed. In projects that use Anaconda, it is also common to see a file named requirements.txt listing all the packages that one has to install.

First, let's import all the necessary modules using the import function. For this exercise, we will be mainly using pandas but you will also gain some experience with numpy, seaborn, and matplotlib. 

To learn more about these packages, you can read through the documentation:  
https://pandas.pydata.org/  
https://numpy.org/  
https://seaborn.pydata.org/  
https://matplotlib.org/

```python
# Let's import the relevant Python packages here
# Feel free to import any other packages for this project

#Data Wrangling
import numpy as np
import pandas as pd

#Plotting
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns

%matplotlib inline
```

## Data
In the following exercise, we will familiarize ourselves with Python by studying the College dataset, which can be found in the file College.csv. This dataset contains the following variables from 777 different universities and colleges in the US. Just try to skim through the table to get an idea of all the datapoints we are provided with.

| Column | Description |
|:-|:-|
|Private | Public/private indicator|
|Apps | Number of applications received|
|Accept | Number of applicants accepted|
|Enroll | Number of new students enrolled|
|Top10perc | New students from top 10\% of high school class|
|Top25perc | New students from top 25\% of high school class|
|F.Undergrad | Number of full-time undergraduates|
|P.Undergrad | Number of part-time undergraduates|
|Outstate | Out-of-state tuition|
|Room.Board | Room and board costs|
|Books | Estimated book costs|
|Personal | Estimated personal spending|
|PhD | Percent of faculty with Ph.D.’s|
|Terminal | Percent of faculty with terminal degree|
|S.F.Ratio | Student/faculty ratio|
|Perc.alumni | Percent of alumni who donate|
|Expend | Instructional expenditure per student|
|Grad.Rate | Graduation rate|

Load the College dataset using pandas

```python
college = pd.read_csv("College.csv")

# Renaming a few columns for readability 
college = college.rename(columns = {'Grad.Rate':'Grad_Rate',
                              'S.F.Ratio': 'S_F_Ratio',
                              'perc.alumni':'perc_alumni',
                              'Room.Board':'Room_Board',
                              'F.Undergrad':'F_Undergrad',
                              'P.Undergrad':'P_Undergrad'})

```

1. Use the head() function to view a few rows in the data.

<!-- grade_id: cell-solution-1, points: 0, grade: false, solution: true -->
```python
# Try to use head() on 'college' and observe the output

# your code here
raise NotImplementedError

```

Notice that there is a column ‘Names’ of each university’s name. As we don’t want to use these names as predictors, they are natural candidates to index our data. Indexing simply means using these college names to represent the corresponding row in our data, instead of serial numbers. We can do this using the following function.

```python
college.set_index('Names', inplace = True)
```

Let's use the head() function again. You should now see that the indices have been replaced with the name of each university in the data set. This means that Python has given each row a name corresponding to the appropriate university. Python will not try to perform calculations on the row names.

```python
college.head()
```

2. Use the info() function to check and produce a numerical summary of your variables.

<!-- grade_id: cell-solution-4, points: 0, grade: false, solution: true -->
```python
# Try to use info() for 'college'and skim through the result

# your code here
raise NotImplementedError

```

3. Examine if there are any duplicates and drop them if needed. Drop them "in place".

Hint 1: It is perfectly fine to observe no duplicates in the dataset, but it is good practice to check anyhow. Before dropping duplicates, it is also good practice to check whether the duplicates contained any different data, since data entry errors frequently occur.

Hint 2: Check for "inplace" argument in the drop duplicates documentation 

<!-- grade_id: cell-solution-5, points: 0, grade: false, solution: true -->
```python
# Modify 'college' inplace to drop duplicates inplace

# your code here
raise NotImplementedError

college.info() #check again
```

4. Replace any missing values in the 'Apps' column with 0. This dataframe will henceforth be our original dataframe. 

Hint: Na means no value, NaN means Not a Number. It is perfectly fine if there are no Na's in our dataset, but we should always check just in case.

<!-- grade_id: cell-solution-6, points: 0, grade: false, solution: true -->
```python
# Modify 'college' inplace to replace any missing values in the 'Apps' column with 0

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-032bedd4d7fd0ad8, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert college[college['Apps'] == 0].empty, "Filtered colleges for those with 0 apps dataframe should be empty.."
```

5. Find the college with the least out-of-state tuition and name this variable `college_least_tuition`. The variable should return the name of a college, not its tuition. 

<!-- grade_id: cell-solution-7, points: 0, grade: false, solution: true -->
```python
college_least_tuition = None # Save your solution in this variable

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-8108a4e9b615b0f8, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert college_least_tuition == "Brigham Young University at Provo", "college_least_tuition is wrong, try again.."
```

6. From the original dataframe, select the ‘PhD’ column and name this dataframe phd_column. Find the length of this column and name this variable `phd_column_length`.

Hint: Make sure to use double brackets to select a Pandas dataframe and single brackets to select a Pandas series. You will notice the difference in formatting: Pandas dataframes are neatly formatted, while Pandas series are not formatted. A Pandas dataframe gives extra functionality compared to a Pandas series, such as appending other dataframes and selecting multiple columns. A Pandas series is essentially a Numpy column.

<!-- grade_id: cell-solution-8, points: 0, grade: false, solution: true -->
```python
# Double bracket for dataframe

phd_column = None # Save your solution in this variable

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-6ff0c25e701ce48c, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

phd_column_length = len(phd_column)
assert phd_column_length==777, "Length of phd_column is wrong, try again.."
```

7. (a) From the original dataframe, select both the ‘Private’ and ‘Top10perc’ columns, and slice them such that only the rows with index 15 and 16 remain. Name this dataframe private_top10. 
   
   (b) From this dataframe, find the length of the filtered ‘Private’ column and name this variable `private_column_length`.

<!-- grade_id: cell-solution-10, points: 0, grade: false, solution: true -->
```python
# Double bracket to obtain multiple columns

private_top10 = None # Save your solution in this variable

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-solution-11, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

private_column_length = len(private_top10[['Private']])
assert private_column_length == 2, "Length of private_top10 is wrong, try again.."
```

From the original dataframe, select the row that only contains data about the “University of Pennsylvania”. Note that many other colleges share the same name, but there is only one unique University of Pennsylvania. Name this dataframe penn.

```python
penn = college.loc[['University of Pennsylvania']]
penn
```

8. From the original dataframe, select the rows that contain all colleges with the name “Penn” included. The “P” in Penn should be capitalized. Name this dataframe `many_penns`.

<!-- grade_id: cell-solution-13, points: 0, grade: false, solution: true -->
```python
many_penns = None # Save your solution in this variable

# your code here
raise NotImplementedError

many_penns
```

<!-- grade_id: cell-adbf8518298036c7, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert many_penns.index.str.contains("Penn").all(), "Values in many_penns is wrong, try again.."
assert len(many_penns) == 8, "Length of many_penns is wrong, try again.."
```

9. Use the `seaborn` `pairplot()` function on the original dataframe with all universities to produce a scatterplot matrix of the ﬁrst ten columns of the data.

<!-- grade_id: cell-solution-14, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

10. Use the `seaborn` `boxplot()` function on the original dataframe with all universities to produce side-by-side box plots of ‘Outstate’ versus ‘Private’. Find the average out of state tuition at private universities and name this variable`private`. Find the average out of state tuition at public universities and name this variable `public`. An approximate solution for both is fine!

<!-- grade_id: cell-solution-15, points: 0, grade: false, solution: true -->
```python
private = None # Save your solution in this variable
public = None # Save your solution in this variable
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-test-15, points: 10, grade: true, solution: false -->
```python
# ##########################
# ### TEST YOUR SOLUTION ###
# ##########################
assert np.allclose(private, 11000, 1000), "Value for private is wrong, try again.."
assert np.allclose(public, 6000, 1000), "Value for public is wrong, try again.."
```

11. Create a new qualitative variable named `large_university` by binning the ‘Enroll’ column. We are going to divide universities into two groups based on whether the number of new students enrolled exceeds the average (mean) of all new students enrolled.  Hint: The variable `large_university` is an indicator variable (a Pandas series that represents a series of booleans, a list is also acceptable). In Python, this is called a "mask," which you can later manipulate to slice your data.

<!-- grade_id: cell-solution-16, points: 0, grade: false, solution: true -->
```python
large_university = None # Save your solution in this variable
# your code here
raise NotImplementedError
```

12. Create a dataframe called `large_universities` that only includes large universities.

<!-- grade_id: cell-solution-17, points: 0, grade: false, solution: true -->
```python
large_univsersities = None # Save your solution in this variable
# your code here
raise NotImplementedError
large_universities.head()
```

13. Use the `pandas` `describe(include = ‘all’)` function on large_universities to produce a numerical summary of each column. Within this dataframe, find the $75^{th}$ percentile for the column ‘Enroll’ and name this variable `enroll`.

<!-- grade_id: cell-solution-18, points: 0, grade: false, solution: true -->
```python
enroll = None # Save your solution in this variable
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-test-18, points: 10, grade: true, solution: false -->
```python
# ##########################
# ### TEST YOUR SOLUTION ###
# ##########################
assert np.allclose(enroll, 2408.0, .01), "Value for enroll is wrong, try again.."
```

14. Use `matplotlib.pyplot.hist()` or `seaborn.distplot()` to create a histogram to examine the number of applications received with a bin size of 200. Use your histogram to answer the following question. Is it true that more universities receive 0 to 10,000 applications than 10,000 to 20,000 applications? Store your answer in a variable named `applications`.

<!-- grade_id: cell-solution-19, points: 0, grade: false, solution: true -->
```python
applications = None # Save your solution in this variable
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-test-19, points: 10, grade: true, solution: false -->
```python
# ##########################
# ### TEST YOUR SOLUTION ###
# ##########################
assert applications == True, "Value for applications is wrong, try again.."
```

15. We are interested in the acceptance rate of these universities. To calculate the acceptance rate, we take the ‘Accept’ column and divide by the ‘Apps’ column. Add this column to a copied version of the original dataframe. Remember to copy the dataframe first before performing feature transformation, as we do not wish to alter or inadvertently create new columns in our original data. Name your new column `acceptance_rate` and use your new dataframe to answer the following questions. Which college has the lowest acceptance rate? Which college has the highest acceptance rate? Store your answers in variables named `lowest` and `highest` respectively. If there is a tie, you can pick any university.

<!-- grade_id: cell-solution-20, points: 0, grade: false, solution: true -->
```python
lowest = None # Save your solution in this variable
highest = None # Save your solution in this variable
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-test-20, points: 20, grade: true, solution: false -->
```python
# ##########################
# ### TEST YOUR SOLUTION ###
# ##########################
assert lowest == "Princeton University", "Value for highest is wrong, try again.."
assert highest in ['Emporia State University', 'Maryville State University', 'MidAmerica Nazarene College', 'Southwest Baptist University', 'University of Wisconsin-Superior', 'Wayne State College'], "Value for lowest is wrong, try again.."
```

```python

```